# Homework 3: Files, Objects, and Functional Programming (26 points)

name: Yiyang Zhang

email: zhyiyang@umich.edu

This homework assignment took me 15 hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).

I did not discuss this homework with anyone.

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. 


Before submitting, please make sure to __Cells->Run All__ executes without errors; errors in your code translate directly to point deductions. 
In general, you don't need to do explicitly raise errors (e.g. with the ```raise``` function) if we don't ask you to in the problem statement.
However, even in cases where we ask you to check for errors, your submission should not contain any examples of your functions actually raising those errors.

Note that many parts of this homework where you are expected to type in code will have ```NotImplementedError()``` as a placeholder. You need to delete this function and replace it with your own code.

## Importing libaries

Unless specifically prohibited, you can import any Python library that comes included with Python (e.g. ```math```, ```datetime```) or that has been used in the lecture or homework prompts. You do not need any other libraries to complete the homework. Note that the graders will not install any additional libraries to grade an assignment.


<!-- ## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking: -->

## Homework tips 

1. **Start early!** If you run into trouble installing things or importing packages, it’s
best to find those problems well in advance, not the night before your assignment is
due when we cannot help you!

2. **Make sure you back up your work!** At a minimum, do your work in a Dropbox
folder. Better yet, use git, which is well worth your time and effort to learn.

3. **Be careful to follow directions!** Remember that Python is case sensitive. If
you are ask you to define a function called my_function and you define a function
called My_Function, you will not receive full credit. You may want to copy-paste
the function names below to make sure that the functions in your notebook match.

## Error checking

You do not need to do error checking (raising errors, etc.) in your code unless we explicitly ask you to so in a problem.

## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking:


## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. You will notice some `read-only` cells in the assignment that contain `assert` statements. These are tests that your code must pass for your solution to be correct. If any of the tests fail, you will get an python error and not get points for that question. 

**Note:** The tests shown not are not comprehensive; additional tests will be used at grading time. You are encouraged to read the problem carefully and verify your code covers all possible cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.


## Problem 1: Counting Word Bigrams (6 points) 
In your previous homework, you wrote a function for counting character bigrams. Now, let’s write a function for counting word bigrams. That is, for each pair of words, say, cat and dog, we want to count how many times the word “cat” occurred immediately before the word “dog”. We will represent this bigram by a tuple, `(’cat’, ’dog’)`. 

So, as an example, the fragment of poem,

```
Half a league, half a league, Half a league onward,
All in the valley of Death Rode the six hundred.
```

includes the bigrams `(’half’, ’a’)` and `(’a’, ’league’)` both three times, the bigram `(’league’, ’half’)` appears twice, while the bigram `(’in’, ’the’)` appears only once.

**Note:** For our purposes, we will ignore all spaces, newlines, punctuation and capitalization in our counting.

### Part 1 (2 points) 
Write a function `count_bigrams_in_file` that takes a ﬁlename `file` as its only argument. Your function should read from the given ﬁle, and return a dictionary whose keys are bigrams (given in the tuple form above), and values are the counts for those bigrams. 

Again, your function should ignore punctuation, spaces, newlines and capitalization.

#### Hints

- You will ﬁnd the Python function `str.strip()`, along with the string constants deﬁned in the string documentation (https://docs.python.org/3/library/string.html), useful in removing punctuation. You are free to use the `re` library for this problem but our solution does not use it. 
- Be careful to check that your function handles newlines correctly. For example, in the poem above, one of the `(’league’, ’half’)` bigrams spans a newline, but should be counted nonetheless. 
- For this function to be executed at a reasonable speed, you should attempt to remove all punctuation from the file in the text __before__ splitting it into words. For reference, our solution only uses one for loop over all the words in the given file's text. 
- Be careful that your function does not accidentally count the empty string as a word (this is a common bug if you aren’t careful about splitting the input text). Solutions that merely delete “bad” keys from the dictionary at the end will not receive full credit, as all edge cases can handled by correctly splitting the input.
- Please replace instances of '--' with a space ' '. So that e.g. 'Bob--David' would be parsed as 'Bob David' (two separate words). Any other punctuation should be removed and treated the same, so for example 'Bob/David' would be parsed as 'BobDavid'. This includes punctuation inside words so e.g. "they're" would become "theyre"

<!-- grade_id: cell-00acb788c373a75a, points: 0, grade: false, solution: true -->
```python
import string
import re

def count_bigrams_in_file(file):
    # YOUR CODE HERE
    
    # begin answer
    f = open(file).read()
    fhand = f
    fhand = fhand.lower()
    fhand = fhand.replace("\n","")
    fhand = fhand.replace("\r","")
    fhand = fhand.replace("--"," ")
    fhand = re.sub(r'[^\w\s]', '', fhand) 
    
    fhand_w = fhand.split(" ")
    
    dic = {}
    for i in range(len(fhand_w)-1):
        if (fhand_w[i].strip(),fhand_w[i+1].strip()) in dic.keys():
            dic[(fhand_w[i].strip(),fhand_w[i+1].strip())] = dic[(fhand_w[i].strip(),fhand_w[i+1].strip())] + 1
        else:
            dic[(fhand_w[i].strip(),fhand_w[i+1].strip())] = 1
    
    return dic
    # end answer
```

<!-- grade_id: cell-64a1f6983f806530, points: 2, grade: true, solution: false -->
```python
"""Test simple cases"""
sol = {('a', 'league'): 3, ('half', 'a'): 3,
       ('league', 'half'): 2, ('league', 'onward'): 1}
with open('./test.txt', 'w') as f:
    f.write('Half a league, half a league,\n Half a league onward.')
assert count_bigrams_in_file('test.txt') == sol
```

### Part 2 (1 point)  
The ﬁle `WandP.txt`, which was included along with this jupyter notebook in `hw3.zip`, is an `ASCII` copy of all of Tolstoi’s novel War and Peace. Run your function on this ﬁle, and pickle the resulting dictionary in a ﬁle called `mb.bigrams.pickle`.

<!-- grade_id: cell-957e399ee02c11d9, points: 0, grade: false, solution: true -->
```python
"""Code for creating 'mb.bigrams.pickle' here."""
import pickle
# YOUR CODE HERE

# begin answer
file = "WandP.txt"
mydict = count_bigrams_in_file(file)
pickle.dump(mydict, open("mb.bigrams.pickle", "wb"))

# end answer
```

<!-- grade_id: cell-43238c586b58202d, points: 1, grade: true, solution: false -->
```python
"""Check that your pickled file is exists and is correct."""
import os
import pickle

assert os.path.isfile('mb.bigrams.pickle')
```

### Part 3 (2 points)

We say that word `A` is _collocated_ with word `B` in a text if words `A` and `B` occur immediately one after another (in either order). That is, words `A` and `B` are collocated if and only if either of the tuples `(A, B)` or `(B, A)` are present in the text. 

Write a function `collocations` that takes a ﬁlename `file` as its only argument and returns a dictionary. Your function should read from the given ﬁle and return a dictionary whose keys are all the strings appearing in the ﬁle and the value of word `A` is a Python set containing all the words collocated with `A`. 

Follow the same conventions as Part 1/2 in ignoring case and stripping away all spaces, newlines and punctuation.

Again using the poem fragment above as an example, the string `’league’` should appear as a key, and should have as its value the set `{’a’, ’half’, ’onward’}`, while the string `’in’` should have the set `{’all’, ’the’}` as its value. 

#### Hints
- We didn’t discuss Python sets in lecture, because they are essentially just dictionaries without values. See the documentation at https://docs.python.org/3/tutorial/datastructures.html#sets for more information.
- For this function to be executed at a reasonable speed, you should attempt to remove all punctuation from the file in the text __before__ splitting it into words. For reference, our solution only uses one for loop over all the words in the given file's text. 

<!-- grade_id: cell-9a8dbd9611c8eb06, points: 0, grade: false, solution: true -->
```python
def collocations(file):
    # YOUR CODE HERE
    
    # begin answer
    f = open(file).read()
    fhand = f
    fhand = fhand.lower()
    fhand = fhand.replace("\n","")
    fhand = fhand.replace("\r","")
    fhand = fhand.replace("--"," ")
    fhand = re.sub(r'[^\w\s]', '', fhand) 
    
    fhand_w = fhand.split(" ")
    
    dic1 = {}
    dic1[fhand_w[0]] = set([fhand_w[1]])
    dic1[fhand_w[len(fhand_w)-1]] = set([fhand_w[len(fhand_w)-2]])
    for i in range(1,len(fhand_w)-1):
        if fhand_w[i] in dic1.keys():
            dic1[fhand_w[i]].add(fhand_w[i-1])
            dic1[fhand_w[i]].add(fhand_w[i+1])
        else:
            dic1[fhand_w[i]] = {fhand_w[i-1],fhand_w[i+1]}
    return dic1
    # end answer
```

<!-- grade_id: cell-9fa29d301a3567b5, points: 2, grade: true, solution: false -->
```python
"""Test it works on simple case."""
with open('test.txt', 'w') as f:
    f.write('Half a league, half a league,\n Half a league onward.')
assert collocations('test.txt') == {
                             'half': {'league', 'a'},
                             'a': {'half', 'league'},
                             'league': {'onward', 'half', 'a'},
                             'onward': {'league'}
                              }
```

### Part 4 (1 point) 
Run your function on the ﬁle `WandP.txt` and pickle the resulting dictionary in a ﬁle called `mb.colloc.pickle.`

<!-- grade_id: cell-49a3dcf4ea8de03f, points: 0, grade: false, solution: true -->
```python
"""Code for creating 'mb.colloc.pickle'."""
import pickle
# YOUR CODE HERE

# begin answer
file = "WandP.txt"
mydict = collocations(file)
pickle.dump(mydict, open("mb.colloc.pickle", "wb"))
# end answer
```

<!-- grade_id: cell-04c1ab0c9d863a66, points: 1, grade: true, solution: false -->
```python
"""Check that your pickled file exists and is correct"""
import os
import pickle
assert os.path.exists('mb.colloc.pickle')
```

## Problem 2: More Fun with Vectors (12 points)

In this exercise, we’ll encounter our old friend the vector yet again, this time taking an object-oriented approach.

**Note:** Before starting this problem, you are strongly encouraged to read _Chapter 14: Python Objects_ in the textbook: https://www.py4e.com/html3/14-objects

<!-- grade_id: cell-5b18f2366afe4b38, points: 0, grade: false, solution: true -->
```python
"""Include all the code for the Vector class here."""
import math
import numpy as np

class Vector:
    # YOUR CODE HERE
    
    # begin answer
    def __init__(self, dimension = 0, entries = [] or ()):
        if dimension < 0:
            raise ValueError
        if type(entries) != list and type(entries) != tuple:
            raise TypeError("Entries should be list or tuple")
        if type(dimension) != float and type(dimension) != int:
            raise TypeError("Dimensions should be int or float")
        if dimension != len(entries):
            raise ValueError("Dimensions are not matching")
        self.dimension = dimension
        self.entries = entries
        
    def __ne__(self,other):
        return tuple(self.entries) != tuple(other.entries)
    
    def __le__(self,other):
        return tuple(self.entries) <= tuple(other.entries)
    
    def __ge__(self,other):
        return tuple(self.entries) >= tuple(other.entries)
    
    def __eq__(self,other):
        return tuple(self.entries) == tuple(other.entries)
    
    def __mul__(self,other):
        if type(other) != int:
            if len(self.entries) != len(other.entries):
                raise ValueError("The dimesions are not match")
            ans = []
            for i in range(len(self.entries)):
                ans.append(self.entries[i]*other.entries[i])
            return Vector(self.dimension,tuple(ans))
        else:
            ans = []
            for i in range(len(self.entries)):
                ans.append(self.entries[i]*other)
            return Vector(self.dimension,tuple(ans))
            
    def __rmul__(self,other):
        ans = []
        for i in range(len(self.entries)):
            ans.append(self.entries[i]*other)
        return Vector(self.dimension,tuple(ans))
                
    def dot(self,other):
        if len(self.entries) != len(other.entries):
            return ValueError("The dimesions are not match")
        return np.dot(tuple(self.entries),tuple(other.entries))
            
    def norm(self,p):
        normvalue = 0
        if p == 0:
            for i in range(len(self.entries)):
                if self.entries[i] != 0:
                    normvalue += 1
            return normvalue
        elif p > 0 and p < math.inf:
            for i in range(len(self.entries)):
                normvalue += self.entries[i]**p
            normvalue = normvalue**(1/p)
            return normvalue
        elif p == math.inf:
            for i in range(len(self.entries)):
                if self.entries[i] > normvalue:
                    normvalue = self.entries[i]
            return normvalue

    # end answer
```

### Part 1 (5 points)
Deﬁne a class `Vector`. The initializer for your class should take the dimension as its ﬁrst argument and a `list` or `tuple` of numbers (`int`s or `ﬂoat`s), representing the vector’s entries, as its second argument (**Note:** your initializer should work correctly given either a list or a tuple of numbers). The initializer should create a vector with dimension $0$ and no entries (an empty tuple `tuple()`) if the user does not provide the corresponding arguments.

Your class vector should have two __instance__ attributes: `dimension` and `entries`. The first containing the dimension of the vector as an `int` and the second containing the entries of the vector as a `tuple`. 

The initializer should raise an appropriate error in the case where the dimension is invalid (i.e., wrong type or a negative number), and should also raise an error in the event that the dimension and the number of supplied entries disagree.

**Note:** Please raise `ValueError` or `TypeError` for your error checking in this problem. Raising either will work. 

<!-- grade_id: cell-d84d85fc86863cd5, points: 1, grade: true, solution: false -->
```python
"""Check initialization works on sensible cases"""
v = Vector(4, (1, 2, 3, 4))
v = Vector(4, [1, 2, 3, 4])
v = Vector(1, [1])
v = Vector(2, [2, 3])
v = Vector(1, (1,))
v = Vector(0, [])
v = Vector(0)
assert True
```

<!-- grade_id: cell-ee5b2851ac09f0d4, points: 1, grade: true, solution: false -->
```python
"""Vector has correct attributes"""
v = Vector(4, (1, 2, 3, 4))
assert isinstance(v, Vector)
assert 'entries' in v.__dict__.keys()
assert 'dimension' in v.__dict__.keys()
assert type(v.dimension) == int
assert type(v.entries) == tuple
```

<!-- grade_id: cell-556867b4cb67ffaf, points: 1, grade: true, solution: false -->
```python
"""Check you can get entries from vector and they are correct when retrieved"""
v = Vector(3, [1, 2, 3])
assert tuple(v.entries) == (1, 2, 3)
```

<!-- grade_id: cell-9eac4809b5a757a5, points: 2, grade: true, solution: false -->
```python
"""Check exceptions are handled correctly"""
try:
    Vector(5, (1, 2, 3, 4))
    assert False
except (ValueError, TypeError):
    assert True

try:
    Vector(-4, (1, 2, 3, 4))
    assert False
except (ValueError, TypeError):
    assert True

try:
    Vector('a', [1, 2, 3, 4])
    assert False
except (ValueError, TypeError):
    assert True

try:
    Vector(1, [])
    assert False
except (ValueError, TypeError):
    assert True
```

### Part 2 (2 point)
Implement the necessary operator(s) to support comparison (equality, less than, less or equal to, greater than, etc) of `Vector` objects. We will say that two `Vector` objects are equivalent if they have the same coordinates. Otherwise, comparison should be analogous to `tuple`s in Python, so that comparison is done on the ﬁrst coordinate ﬁrst, then the second coordinate, then the third, and so on. So, for example, the two-dimensional vector `(2, 4)` is ordered before (less than) `(2, 5)`. 

You can assume you are always comparing vectors of the same dimension. 

<!-- grade_id: cell-3a1c2f17cf2a4779, points: 1, grade: true, solution: false -->
```python
assert Vector(3, [2, 1, 1]) != Vector(3, [1, 3, 3])
assert Vector(3, [2, 1, 1]) >= Vector(3, [1, 3, 3])
assert Vector(3, [1, 1, 1]) == Vector(3, [1, 1, 1])
```

<!-- grade_id: cell-de890d7da24aa5f0, points: 1, grade: true, solution: false -->
```python
assert Vector(3, [1, 1, 1]) != Vector(3, [1, 1, 3])
assert Vector(3, [1, 3, 3]) <= Vector(3, [2, 1, 1])
```

### Part 3 (1 point) 
Implement a method `Vector.dot` that takes a single Vector as its argument and returns the inner product of the caller with the given `Vector` object. You can assume the input is always a vector of the same dimension as the original one.

<!-- grade_id: cell-2e45b3413aad5e1c, points: 1, grade: true, solution: false -->
```python
"""Check your Vector.dot() method works correctly"""
assert Vector(3, (1, 2, 3)).dot(Vector(3, (1, 2, 2))) == 11.0
```

### Part 4 (2 points)
We would also like our Vector class to support scalar multiplication. Left- or right multiplication by a scalar, e.g., `2*v` or `v*2`, where `v` is a Vector object, should result in a new `Vector` object with its entries all scaled by the given scalar. We will also follow `R` and `numpy` (which you will learn in a few weeks), and use `*` to denote entrywise vector-vector multiplication, so that for `Vector` objects `v` and `w`, `v*w` results in a new `Vector` object, with the i-th entry of `v*w` equal to the i-th entry of `v` multiplied by the i-th entry of `w`. Implement the appropriate operators to support this multiplication operation. 

Many languages have a convention for dealing with multiplication of vectors that diﬀer in their dimension, but we will not handle this case. Your method should raise a `ValueError` in the event that `v` and `w` disagree in their dimensions.

<!-- grade_id: cell-a5f2178bd7c9f252, points: 1, grade: true, solution: false -->
```python
assert 2 * Vector(3, (1, 2, 3)) == Vector(3, (2, 4, 6))
assert Vector(3, (1, 2, 3)) * 2 == Vector(3, (2, 4, 6))
assert Vector(3, (1, 2, 3)) * Vector(3, (1, 2, 3)) == Vector(3, (1, 4, 9))
```

<!-- grade_id: cell-cd6422e310c07ed2, points: 1, grade: true, solution: false -->
```python
try:
    Vector(3, (1, 2, 3)) * Vector(2, (1, 2))
    assert False
except ValueError:
    assert True
```

### Part 5 (2 points)
For a real number $0 \leq p \leq \infty$, and a vector $v \in \mathbb{R}^{d}$, the $p$-norm of $v$, written $\| v \|_{p}$, is given by: 

$$ 
\| v \|_p = \begin{cases}
                        \sum_{i=1}^d 1_{v_i \neq 0} &\mbox{ if } p = 0\\
                        (\sum_{i=1}^d |v_i|^p )^{1/p} &\mbox{ if }
                                                        0 < p < \infty, \\
                        \max_{i=1,2,\dots,d} |v_i| &\mbox{ if } p=\infty
                        \end{cases}
$$

where $v_i$ is the $i$-th component of the vector and the notation $1_{v_i \neq 0}$ indicates: 

$$ 
1_{v_i \neq 0} = \begin{cases}
                        1 &\mbox{ if } v_i \neq 0\\
                        0 &\mbox{ if } v_i = 0
                        \end{cases}. 
$$


Implement a method `Vector.norm` that takes a single int or ﬂoat `p` as an argument and returns the $p$-norm of the calling `Vector` object. You can assume the input $p$ is always a non-negative integer or float.

`Hint:` see https://docs.python.org/3/library/functions.html# float for documentation on representing positive inﬁnity in Python.

<!-- grade_id: cell-2af28a606f1af266, points: 1, grade: true, solution: false -->
```python
assert Vector(3, (1, 2, 3)).norm(1) == 6.0
assert Vector(2, (1, 2)).norm(2) == math.sqrt(1**2 + 2**2)
```

<!-- grade_id: cell-a8f312c44b5334d0, points: 1, grade: true, solution: false -->
```python
assert Vector(5, (3, 2, 1, 10, 5)).norm(math.inf) == 10
assert Vector(5, (3, 2, 1, 10, 5)).norm(0) == 5
```

##  Problem 3: Iterators and Generators (5 points) 
In this exercise, you’ll get some practice working with iterators and generators. 

**Note:** in this problem, the word enumerate is meant in the sense of returning elements, not in
the sense of the Python function enumerate. So, if I say that "an iterator enumerates a
sequence $a_{0}, a_{1}, a_{2}, ...$", I mean that these are the elements that it returns upon calls to
the `__next__` method, not that it returns pairs $(i,a_{i})$ like the `enumerate` function.

### Part 1 (1 point) 
Deﬁne a class `Fibo` of iterators that enumerate the Fibonacci numbers. For the purposes of this problem, the Fibonacci sequence begins $0, 1, 1, 2, 3, ...$ with the $n$-th Fibonacci number $F_{n}$ given by the recursive formula $F_{n} = F_{n-1} + F_{n-2}$  where $F_{0} = 0$ and $F_{1} = 1$. 

Your solution should not make use of any function aside from addition (i.e., you should not need to use the function `fibo()` deﬁned in lecture a few weeks ago). Your class should support, at a minimum, an initialization method, a `__iter__` method (so that we can get an iterator) and a `__next__` method. 

**Note:** there is an especially simple solution to this problem that can be expressed in just a few lines using tuple assignment.

<!-- grade_id: cell-bbaefcc15e4a7081, points: 0, grade: false, solution: true -->
```python
class Fibo():
    # YOUR CODE HERE
    
    # begin answer
    def __init__(self):
        self.a = 0
        self.b = 1
    def __iter__(self):
        return self
    def __next__(self):
        fib = self.a
        self.a, self.b = self.b, self.a + self.b
        return fib
    # end answer
```

<!-- grade_id: cell-e98351d056391b0a, points: 1, grade: true, solution: false -->
```python
ans1 = Fibo()
assert [next(ans1) for _ in range(10)] == [0, 1, 1, 2, 3, 5, 8, 13, 21, 34]
```

### Part 2 (2 points) 

Deﬁne a generator `primes` that enumerates the prime numbers. Recall that a prime number is any integer $p > 1$ whose only divisors are $p$ and $1$. 

**Note:** you may use the function `is_prime` that we deﬁned in class (or something similar to it) and receive full credit. But there is a more graceful solution that avoids declaring a separate function or method for directly checking primality. 

**Hint:** consider a pattern similar to the one seen in lecture using the any and/or all functions.

<!-- grade_id: cell-6075e6fd4a6029d8, points: 0, grade: false, solution: true -->
```python
import math


def primes():
    # YOUR CODE HERE
    
    # begin answer
    D = {}            
    q = 2

    while True:
        if q not in D:
            yield q
            D[q * q] = [q]
        else:
            for p in D[q]:
                D.setdefault(p + q, []).append(p)
            del D[q]
        q += 1
        
    # end answer
```

<!-- grade_id: cell-39a0a188d4945486, points: 1, grade: true, solution: false -->
```python
ans = primes()
assert [next(ans) for _ in range(10)] == [2, 3, 5, 7, 11, 13, 17, 19, 23, 29]
```

<!-- grade_id: cell-b51f7843f8591a1f, points: 1, grade: true, solution: false -->
```python
ans = primes()
assert [next(ans) for _ in range(1)] == [2]
ans = primes()
assert [next(ans) for _ in range(0)] == []
```

### Part 3 (2 points) 
This one is good practice for coding interview questions. The _Ulam numbers_ are a sequence $u_1 , u_2 , u_3, ...$ of positive integers, deﬁned in the following way: $u_1 = 1$, and $u_2 = 2$. For all $n > 2$, $u_n$ is the smallest integer that is expressible as a sum of two distinct terms from earlier in the sequence in exactly one way. See the Examples section of the Wikipedia page for an illustration: https://en.wikipedia.org/wiki/Ulam_number. 

Deﬁne a generator `ulam` that enumerates the Ulam numbers. 

**Hint:** it will be helpful to try and break this problem into smaller, simpler subproblems. In particular, you may ﬁnd it helpful to write a function that takes a list of integers `t` and one additional integer `u`, and determines whether or not `u` is expressible as a sum of two distinct elements of `t` in exactly one way.

<!-- grade_id: cell-ae05ce7df31047af, points: 0, grade: false, solution: true -->
```python
def is_ulam(t, n):
    # YOUR CODE HERE
    # begin answer
    length = len(t)
    if n <= t[length - 1]:
        return False
    sum_lst = []
    remove_lst = []
    sum_lst.append(1)
    sum_lst.append(2)
    for i in range(length - 1):
        for j in range(i+1, length):
            curr_sum = t[i] + t[j]
            if curr_sum not in sum_lst:
                if curr_sum not in remove_lst:
                    sum_lst.append(curr_sum)
            else:
                if curr_sum == n:
                    return False
                sum_lst.remove(curr_sum)
                remove_lst.append(curr_sum)
    sum_lst.sort()
    if n == sum_lst[length]:
        return True
    else:
        return False
    # end answer

def ulam():
    # YOUR CODE HERE
    # begin answer
    num = 1
    seq = []
    while True:
        if num == 1 or num == 2:
            yield num
            seq.append(num)
            num += 1
        else:
            if is_ulam(seq,num) == True:
                seq.append(num)
                yield num
            num += 1
    # end answer
```

<!-- grade_id: cell-88e21fdf334740cf, points: 1, grade: true, solution: false -->
```python
ans = ulam()
assert [next(ans) for _ in range(18)] == [1, 2, 3, 4, 6, 8, 11, 13, 16, 18,
                                          26, 28, 36, 38, 47, 48, 53, 57]
```

<!-- grade_id: cell-f869c2e123a88ffc, points: 1, grade: true, solution: false -->
```python
ans = ulam()
assert [next(ans) for _ in range(5)] == [1, 2, 3, 4, 6]
```

##  Problem 4: List Comprehensions and Generator Expressions (3 points)
In this problem you’ll write a few simple list comprehensions and generator expressions.

### Part 1 (1 point) 
Write a list comprehension where each element equals $3^{n}−1$ for $n = 1, 2, 3,..., 20$.

For ease of grading, please assign this list comprehension to a variable called `pow3minus1`.

<!-- grade_id: cell-5453d98b9528fcce, points: 0, grade: false, solution: true -->
```python
"""Write the list comprehension here. Make sure to assign it
to a variable `pow3minus1`"""
# YOUR CODE HERE

# begin answer
pow3minus1 = [3**n - 1 for n in range(1,21)]
# end answer
```

<!-- grade_id: cell-3bb76ac1abf8f11e, points: 1, grade: true, solution: false -->
```python
"""Check your pow3minus1 is correct."""
assert pow3minus1[0:2] == [2, 8]
```

### Part 2 (1 point) 
Write a generator expression for the **pyramid numbers**. The $n$-th pyramid number $P_{n}$ (for $n = 1, 2, ...$) counts the number of spheres in a pyramid with an $n$-by-$n$ based (see https://en.wikipedia.org/wiki/Square_pyramidal_number), and is given by:
$$
P_{n} = \sum_{k=1}^{n}k^2 = \frac{n (n+1) (2n+1)}{6}
$$
For ease of grading, please assign this generator expression to a variable called `pyramid`. 

**Hint:** You may ﬁnd it useful to first deﬁne a generator for the positive integers.

<!-- grade_id: cell-191508c2cdde347d, points: 0, grade: false, solution: true -->
```python
"""Your code for the pyramid generator goes here."""


# YOUR CODE HERE

# begin answer
def pyramid_function():
    (pyramid,n) = (0,1)
    while True:
        (pyramid,n) = (n*(n+1)*(2*n+1)/6,n+1)
        yield pyramid
pyramid = pyramid_function()
# end answer
```

<!-- grade_id: cell-8eb2d2188e8b54d2, points: 1, grade: true, solution: false -->
```python
"""Check your pyramid generator is correct"""
pyramid_list = [next(pyramid) for _ in range(10)]
assert pyramid_list[:5] == [1, 5, 14, 30, 55]
```

### Part 3 (1 point) 
Write a generator expression that enumerates the _octahedral numbers_. The $n$-th octahedral number ($n = 1, 2, ...$) is given by 

$$O_{n} = \frac{n(2n^{2}+1)}{3},$$ 

and counts the number of spheres in an octahedron with $n$ spheres to each edge (see https://en.wikipedia. org/wiki/Octahedral_number). 

For ease of grading, please assign this generator expression to a variable called `octa`. 

**Note:** You can solve this problem any way you see fit. However, a particularly clever solution uses the fact that the $n$-th octahedral number can be expressed as $O_{n} = P_{n} + P_{n-1}$ where $P_{n}$ denotes the $n$-th pyramidal number (which you implemented in the previous subproblem).

<!-- grade_id: cell-94f79e207c8683e2, points: 0, grade: false, solution: true -->
```python
"""Code for `octa` generator here."""


# YOUR CODE HERE

# begin answer
def octa_generator():
    (octa,n) = (1,1)
    while True:
        (octa,n) = (n*(2*n**2+1)/3, n+1)
        yield octa
octa = octa_generator()
# end answer
```

<!-- grade_id: cell-c0ec74ed06e7b6e9, points: 1, grade: true, solution: false -->
```python
"""Test your octa generator is correct"""
octa_list = [next(octa) for _ in range(10)]
assert octa_list[:3] == [1, 6, 19]
```

```python

```

```python

```